[1541024897] SERVICE ALERT: localhost;Disk Space;CRITICAL;SOFT;1;DISK CRITICAL - /run/user/1000/gvfs is not accessible: Permission denied
[1541024957] SERVICE ALERT: localhost;Disk Space;CRITICAL;SOFT;2;DISK CRITICAL - /run/user/1000/gvfs is not accessible: Permission denied
[1541024997] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 22: Connection refused
[1541025017] SERVICE ALERT: localhost;Disk Space;CRITICAL;SOFT;3;DISK CRITICAL - /run/user/1000/gvfs is not accessible: Permission denied
[1541025057] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;2;connect to address 127.0.0.1 and port 22: Connection refused
[1541025077] SERVICE ALERT: localhost;Disk Space;CRITICAL;HARD;4;DISK CRITICAL - /run/user/1000/gvfs is not accessible: Permission denied
[1541025117] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;3;connect to address 127.0.0.1 and port 22: Connection refused
[1541025177] SERVICE ALERT: localhost;SSH;CRITICAL;HARD;4;connect to address 127.0.0.1 and port 22: Connection refused
[1541099692] SERVICE ALERT: localhost;HTTP;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 80: Connection refused
[1541099752] SERVICE ALERT: localhost;HTTP;OK;SOFT;2;HTTP OK: HTTP/1.1 200 OK - 11192 bytes in 0.000 second response time
[1541109692] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 254 processes
[1541109752] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;2;PROCS WARNING: 254 processes
[1541109812] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;3;PROCS WARNING: 254 processes
[1541109872] SERVICE ALERT: localhost;Total Processes;WARNING;HARD;4;PROCS WARNING: 253 processes
[1541111372] SERVICE ALERT: localhost;Total Processes;OK;HARD;4;PROCS OK: 223 processes